## `lab3`—Types, Functions, & Scope

❖ Objectives

-   Understand distinctions in Python's type system.
-   Use indexing to access components of a data structure.
-   Explain the use and composition of functions.
-   Utilize the built-in library of Python functions effectively.
-   Explain blocks and variable scope and understand how it impacts variable use.

### 1. How numbers (and other values) work (Part 2 of ∞)

#### 1.1 Numbers

In the previous lab, you encountered integers and floating-point numbers, the basic numeric data types of Python.  Let's look at these and more *data types*, which are simply *kinds of data* you can describe using Python.  You can determine the type of a number by using the intuitively-named `type` function:

```python
type(1)
```

Output:
```
int
```

```python
type(1.)    # Again, note that `1` is an `int` while `1.` is a `float`.
```

Output:
```
float
```

```python
type(-1)
```

Output:
```
int
```

```python
type(0.)
```

Output:
```
float
```

```python
type(1.564e15)
```

Output:
```
float
```

No surprises there.  *Promotion* is the phenomenon which occurs when (compatible) data types of different kinds are operated on together.  For instance, a `float` plus an `int` is a `float`:

```python
type(1e5+1)
```

Output:
```
float
```

Sometimes functions change the type of a number, even if the result could be represented as the original type (`int`).  (This can occur when intermediate steps in a calculation use `float`s.)  Sometimes they don't.

```python
from math import exp  # include a library function for the exponential
print('exp(0) = ', exp(10))
type(exp(0))
```

Output:
```
exp(0) =  22026.465794806718
float
```

```python
type(abs(-1))
```

Output:
```
int
```

```python
type(abs(-1.))
```

Output:
```
float
```

Python even supports a handful of special values, such as infinity ∞.

```python
float('inf')  #(It's a little cumbersome to use, though.)
```

Output:
```
inf
```

<h4 style="color:#FF8C00">Exercises</h4>

-   What is the type of the result of the following expression?
        
        int(exp(abs(-2)) + abs(-1.5)) // 4.5

<!-- grade_id: expr1, points: 1, grade: true, solution: true -->
```python
# add the type function for this expression and store the result in the variable T
from math import exp
T = int(exp(abs(-2)) + abs(-1.5)) // 4.5  # your code here
type(T)
```

Output:
```
float
```

-   What is the type of the result of the following expression?
        
        abs(int(1.5e4))

<!-- grade_id: expr2, points: 1, grade: true, solution: true -->
```python
# add the type function for this expression and store the result in the variable S
S = abs(int(1.5e4)) # your code here
type(S)
```

Output:
```
int
```

#### 1.2 Strings & Indexing

The fundamental text data type, the string, has already shown up in a few places:
    
    print("I saw this morning morning's minion")

In order to distinguish the human-readable message from the machine command `print`, we surrounded the string with some set of quotes.  *Indexing* refers to the selection of a single letter in a string (more generally, an element of a sequence).  We do this by placing a bracketed number after a string variable:

```python
message = "Sixty zippers were quickly picked from the woven jute bag."
print(message[3])
```

Output:
```
t
```

Recall why `message[3]` returns not `'x'` but `'t'`.  Typically we would consider `'t'` to be the fourth element of the string, but many programming languages—including Python—instead start counting at zero. In this model, the string message is indexed as follows:

| index  | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7 | 8 | ... | 54 | 55 | 56 | 57 |
|--------|---|---|---|---|---|---|---|---|---|-----|----|----|----|----|
| letter | S | i | x | t | y |   | z | i | p | ... | b  | a  | g  | .  |

Aside from starting with zero, indexing is pretty straightforward.

<div class="alert alert-info">
<p>
Programming languages that count from zero are referred to as *zero-indexed*, while those that start from one are *one-indexed*.  The rationale for zero-indexing is that what is actually counted is the *offset* from the beginning of the string—thus `'i'` is *one* away from the beginning, `'x'` is *two* away, etc.  Languages that are “related” to [C](https://en.wikipedia.org/wiki/C_(programming_language) generally adopt this convention.
</p><p>
One argument for one-indexing is that humans (and mathematicians) count from one "naturally".  Languages related to [Fortran](https://en.wikipedia.org/wiki/Fortran) often use one-indexing.
</p><p>
If you'd like to overthink this, please refer to [this Q&A at Stack Overflow](https://programmers.stackexchange.com/questions/110804/why-are-zero-based-arrays-the-norm).
</p>
</div>

```python
message[25]
```

Output:
```
'y'
```

```python
message[1]
```

Output:
```
'i'
```

```python
message[60]
```

Output:
```
IndexError                                Traceback (most recent call last)
<ipython-input-16-1f8a2855e26f> in <module>()
----> 1 message[60]

IndexError: string index out of range
```

(In the last case, we tried to index past the end of the string—**we can't do that!**)

Regular indexing is a fairly natural process for us:  it's just counting off with whole numbers.  What happens if you use a *negative* index?

```python
message[-1]
```

Output:
```
'.'
```

Negative indexing *counts down* from the end of the string.  This is very convenient and saves you the trouble of figuring out the length of the string and then subtracting from that.

```python
message[-2]
```

Output:
```
'g'
```

Some things probe the limits of sanity, though:

```python
message[8.5]
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-19-54811f561653> in <module>()
----> 1 message[8.5]

TypeError: string indices must be integers
```

<h4 style="color:#FF8C00">Exercises</h4>

What is the index of the following letters in this string?

`"You're not fully clean till you're Zestfully clean."`

```python
# definition
string = "You're not fully clean till you're Zestfully clean."
```

- `'n'` (first one)

<!-- grade_id: zest1, points: 1, grade: true, solution: true -->
```python
# your answer here, as a variable ans
ans = string.index('n')
print(ans)
```

Output:
```
7
```

- `'y'` (last one—use a negative index)

<!-- grade_id: zest2, points: 1, grade: true, solution: true -->
```python
# your answer here, as a variable ans
ans = string.rindex('y')
print(ans)
```

Output:
```
43
```

#### 1.3 Slicing

Now we can generalize our notion of *indexing* to *slicing*.  Slicing means using a *range of indices* to take a *slice* of a string.  A range is written `string[left:right]`, where the `left` index is part of the string but the `right` index is not.  (We call this a *half-open* range.)

```python
pangram = 'Sphinx of black quartz: judge my vow.'
```

```python
pangram[0:5]
```

Output:
```
'Sphin'
```

```python
pangram[6:13]
```

Output:
```
' of bla'
```

```python
pangram[0:0]  # Note that the right-hand index is not in the range, so this string is empty.
```

Output:
```
''
```

```python
pangram[0:-1] # Note that -1 is the last index
```

Output:
```
'Sphinx of black quartz: judge my vow'
```

```python
pangram[5:]  # We don't have to specify both sides; if we don't, Python assumes we are going all the way to the end or coming all the way from the beginning.
```

Output:
```
'x of black quartz: judge my vow.'
```

```python
pangram[:10]
```

Output:
```
'Sphinx of '
```

```python
pangram[:]  # Thus this is just a copy of the entire original string.
```

Output:
```
'Sphinx of black quartz: judge my vow.'
```

```python
pangram[4:18:3]  # Take every THIRD character from the fourth space to the eighteenth (exclusive).[left:right:step]
```

Output:
```
'nobcq'
```

```python
pangram[4:-1:3]  # Take every THIRD character from the fourth space to the last (exclusive).[left:right:step]
```

Output:
```
'nobcqr:ueyo'
```

Although strings *are* important and useful, you're still fairly limited in what you're able to do with them.  But many aspects of indexing will carry over to data types you'll meet later, like `list`s and `ndarray`s ($n$-dimensional arrays)—you will use these techniques to manipulate large arrays of numerical data as well.

<h4 style="color:#FF8C00">Exercises</h4>

-   How would you extract the letters `'ssupdla nioiaim'` from the following string?  Write the full expression (*i.e.*, `showoff[???]`) in the answer block below.
        
        showoff = 'sesquipedalian antinominalism'

```python
# definition
showoff = 'sesquipedalian antinominalism'
```

<!-- grade_id: slice, points: 1, grade: true, solution: true -->
```python
# your answer here.
showoff[::2]
```

Output:
```
'ssupdla nioiaim'
```

Bonus exercise: reverse `showoff`

```python
# your answer here. Hint: you can use 'step' flexibly
showoff[::-1]
```

Output:
```
'msilanimonitna nailadepiuqses'
```

### 2. How to reuse your code

When you need to carry out a calculation many times (whether simple or complicated), you don't want to write out the formula over and over again.  Besides being inefficient, if you later need to expand the code or correct a bug, you'd have to look everywhere to make sure you changed every case.  It's far better to write the calculation *once*, as a function, and then call that function every time you need it.

You've seen a handful of functions before:  `type`, `exp`, `abs`.  Let's see what's involved with writing your own.

In mathematics, a function is simply a way of relating an input value (often $x$) to an output value (sometimes $y$ or $f(x)$).  For instance, a parabola may be expressed as

$$y = f(x) = x^2$$

and plotted as

![](./img/parabola.png)

Properly speaking, the resulting curve represents the collection of ordered pairs $(x,y)$.  To plot this, we select a discrete set of points in the domain $x$ and calculate the corresponding set of values $y$.

| $x$  | -1.5 | -1.25  | -1    | -0.75 | -0.5  | -0.25 | 0  | .25    | 0.5  | 0.75   | 1 | 1.25  |  1.5  |
|--------|------|--------|-------|-------|-------|-------|----|--------|------|--------|---|-------|-------|
| $y$ | 2.25 | 1.5625 | 1.0 | 0.5625 | 0.25 | 0.0625 |  0 | 0.0625 | 0.25 | 0.5625 | 1 | 1.5625 | 2.25 |

Of course, the advantage of using a computer is that each step is easier to type than to calculate by hand.  If you had to fill in this table, even using Python it would be tedious if you had to reset the variable `x` and retype the expression for `y` by hand every time.

It is very likely that any equation you will wish to use in your studies and career will be more complicated than $y = x^{2}$.  In order to minimize the amount of time typing code, as well as the likelihood of making a mistake due to repetition, we can use functions to capture the operations involved:

```python
def sqr(x):
    result = x * x
    return result
```

Now instead of typing `x**2` or `x*x` each time you need the formula, you can call the function `sqr(x)` instead.  (In this trivial case, that's actually *more* keystrokes, but the principle stands.)

A more realistic candidate is the quadratic equation (as we saw implemented in lab2), or the equation for an arbitrary line, $y = mx + b$, where $m$ is the slope and $b$ is the intercept.  A function which formats a slope and intercept into a string representing a line may look like this:

```python
def line_eqn(m,b):
    return 'y = %f x + %f'%(m,b)
```

```python
print(line_eqn(1.875,4))
```

Output:
```
y = 1.875000 x + 4.000000
```

#### 2.1 `math` Library Functions

*Modules* are collections of functions which allow you to take advantage of the collective work of the scientific and engineering community when carrying out your calculations or building your models.  One of the strengths of Python is its enormous and rich library of modules, which provide both basic and involved operations far beyond the built-in operators.

Let's consider the `math` library.  To use functions like `sin`, `cos`, `exp`, and `log`, you simply need to `import math`:

```python
import math
```

You still need to specify *which* module contains the functions to use them, but otherwise it's quite straightforward:

$$
\sin\left(\frac{\pi}{4}\right)
$$

```python
math.sin(math.pi/4)
```

Output:
```
0.7071067811865475
```

$$
\left[\text{ceil}\left(4.5\right)\right] !
$$

```python
math.factorial(math.ceil(4.5))
```

Output:
```
120
```

<h4 style="color:#FF8C00">Exercises</h4>

-   Write an expression for the following formula:
    
    $$
    \frac{(m+n)!}{(m+1)!(n-1)!}
    $$
    
    (Note that the factorial refers only to the second term of the denominator, not the entire denominator.)

```python
# these are representative values—don't change these
m = 2
n = 3
```

<!-- grade_id: form, points: 0, grade: false, solution: true -->
```python
# write your expression here
expr = math.factorial(m+n)/(math.factorial(m+1)*math.factorial(n-1)) # your code here
```

<!-- grade_id: form-test, points: 1, grade: true, solution: false -->
```python
# this is a test it should pass
from numpy import isclose
assert isclose(expr, 10.0)
print('Success')
```

Output:
```
Success
```

Some modules are quite large, and you may not want to import everything into memory at once.  In that case, you can import a single function from a module as follows:

```python
from math import asin  # arcsine
```

-   What line of Python code could you write to import `log10` from `math`?  (This is the logarithm base 10, $\log_{10}(x)$.)

<!-- grade_id: import, points: 0, grade: false, solution: true -->
```python
# write your expression here
from math import log10
```

<!-- grade_id: import-test, points: 1, grade: true, solution: false -->
```python
# this is a test it should pass
try:
    log10(0.25)
    print('Success')
except NameError:
    raise AssertionError('log10 not imported yet')
```

Output:
```
Success
```

### 3. Application:  Cardinal Sine Function

The bread and butter of working effectively in Python is writing your own functions.

Consider the *cardinal sine function*, $\text{sinc}(x)$, which is used in signal processing by electrical engineers.

$$\text{sinc}(x) = \frac{\sin(x)}{x}$$

![](https://upload.wikimedia.org/wikipedia/commons/thumb/5/59/Si_sinc.svg/700px-Si_sinc.svg.png)

To compose a function which reproduces this behavior for an arbitrary number, we could simply define:

```python
from math import sin
def sinc(x):
    y = sin(x) / x
    return y
```

This short function simply creates a new variable `y` which contains the result of `sin(x)/x` evaluated for the actual value of `x`.  The `return` statement tells Python what the function gives back—in this case, the value of `y`.

```python
sinc(1.5)
```

Output:
```
0.6649966577360363
```

```python
sinc(3.1415)
```

Output:
```
2.949342341572187e-05
```

Critically, *the name of the variable OUTSIDE of `sinc` is not related to the name of the variable INSIDE `sinc`*.  (We say that they inhabit different *scopes*.)  Consider this:

```python
y = 5.0
sinc(y)
```

Output:
```
-0.1917848549326277
```

<h4 style="color:#FF8C00">Exercises</h4>

-   For what value of input `x` does `sinc` fail?  (You can find this by experiment or inspection.)

<!-- grade_id: divbyzero, points: 1, grade: true, solution: true -->
```python
# your answer here, in the form of a float (e.g., 1.0)
sinc(0.)
```

Output:
```
<ipython-input-52-0273c4c8e1b2> in <module>()
      1 # your answer here, in the form of a float (e.g., 1.0)
----> 2 sinc(0.)

<ipython-input-48-f2b4b0959ef5> in sinc(x)
      1 from math import sin
      2 def sinc(x):
----> 3     y = sin(x) / x
      4     return y

ZeroDivisionError: float division by zero
```

<div class="alert alert-warning">
<p>
The `y` in the scope outside of the function is equal to `5.0`, whereas inside of the function `sinc` the same *value* is referred to as `x`, and `y` means something else.
</p><p>
You can call variables any valid name you like, but it's generally good practice to name it something clear in context.  That is,
    
    <pre><code>def sinc(number_of_calabazas):
        return sin(number_of_calabazas) / number_of_calabazas
</code></pre>
is a little too bohemian to represent good programming style.  **Responsible variable names reflect their use.**
</p>
</div>

### 4. Application:  Atmospheric Pressure

A mercury barometer is frequently used to precisely measure the atmospheric pressure.  The height $h$ of a column of mercury adjusts to a point where its weight plus force due to the mercury vapor pressure $p_{\text{vapor}}$ balances the force due to the external atmospheric pressure $p_{\text{atm}}$.  Mathematically, we write,

$$p_{\text{atm}} = \gamma h + p_{\text{vapor}}$$

where $\gamma$ is the specific weight of mercury.  In SI units, $\gamma = 133 \times 10^{3} \,\text{N}\cdot\text{m}^{-3}$.  The vapor pressure of mercury is typically very small in relation to the other values, so we can safely neglect it ($\left. p_{\text{vapor}} \right|_{T=25 ^{\circ}\text{C}} = 26.13 Pa$)<sup>[[NISTIR 6643](http://www.boulder.nist.gov/div838/SelectedPubs/NISTIR.6643.pdf)]</sup>.  A function which calculates the atmospheric pressure given the height of a column of mercury, then, is:

```python
def p_pa(h_Hg):
    # Calculate and return the atmospheric pressure (in Pa) given the height
    # of a column of mercury (in m).
    gamma = 133e3       # N/m^-3
    P = gamma * h_Hg    # Pa
    return P
```

```python
p_pa(0.598)  # The atmospheric pressure in pascals for 598 mmHg.
```

Output:
```
79534.0
```

```python
p_pa(0.598) / 101325  # The atmospheric pressure in the units of standard atmosphere for 598 mmHg.
```

Output:
```
0.7849395509499136
```

We will now introduce the `input` function, which lets you supply a message and ask the user for a value.  (This value then needs to be converted from a string to a `float`.)

```python
h_str = input('Give the height of the column of mercury in m:')
h = float(h_str) # convert to a floating-point number
print('The result is %f pascals.' % p_pa(h))
```

Output:
```
Give the height of the column of mercury in m:2
The result is 266000.000000 pascals.
```

<h4 style="color:#FF8C00">Exercises</h4>

-   Create a function `p_atm` which accepts the height of a column of mercury and returns the atmospheric pressure in standard atmospheres (demonstrated a couple of cells above this).  (This is equal to the pressure in pascals divided by $101\,325\,\text{Pa}/\text{atm}$.  You may accomplish this by either calling `p_pa` and dividing by 101325 or by doing the calculation directly.)

<!-- grade_id: mercury, points: 0, grade: false, solution: true -->
```python
# define your function here---you can copy and paste from above, or write your own from scratch
def p_atm(h):
    return p_pa(h) / 101325
```

<!-- grade_id: mercury-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test
from numpy import isclose
assert isclose(p_atm(1.00), 1.3126079447322971)
print('Success')
```

Output:
```
Success
```

### 5. Application:  The U.S.D.O.T. Bridge Formula

The United States Department of Transportation, as part of its stewardship of the U.S. interstate highway system, maintains code standards for bridge construction.  As a metric for preserving shipping infrastructure, the DOT uses the bridge formula<sup>[[DOT2014](http://ops.fhwa.dot.gov/freight/sw/brdgcalc/calc_page.htm)]</sup> to limit the legal weight-to-length ratio for freight-laden vehicles crossing a bridge.  In simplified form, the bridge formula is

$$W = 500 \left( \frac{LN}{N-1} + 12N + 36 \right)$$

where

-   $W$ is the overall gross weight on any group of consecutive axles;
-   $L$ is the distance in feet between the outer axles of any group of consecutive axles; and
-   $N$ is the number of axles in the group under consideration.

This prevents short trucks with heavy loads from damaging bridges.

![](http://ops.fhwa.dot.gov/Freight/publications/brdg_frm_wghts/images/fig1.gif)

You will implement this formula as a function and use it to answer the following questions.

<h4 style="color:#FF8C00">Exercises</h4>

-   Complete the function `bridge_formula` by adding lines of code to accomplish both tasks indicated.

<!-- grade_id: dot, points: 0, grade: false, solution: true -->
```python
def bridge_formula(L, N):
    '''
    Calculate bridge weight in accordance with federal law.

    Args:
        L:  distance between groups of consecutive axles, ft
        N:  number of axles in group under consideration
    Returns:
        W:  overall gross weight on any group of consecutive axles, lb
    '''

    # Calculate the bridge formula's weight W using L and N.
    # ... your code here ...
    W = 500*(L*N/(N-1) + 12*N + 36)

    # Return the resulting weight W.
    # ... your code here ...
    return W
```

As a test of your function, the following values are correct:  $L = 51$, $N = 5$, $W = 79,500$.  (This describes a $51'$, 5-axle semi which is legally allowed to carry $79,500 \,\text{lb}$, including freight and truck weight—in practice, U.S.D.O.T. rounds to $80,000$ but we won't worry about that layer of complexity here.)

<!-- grade_id: dot-test1, points: 1, grade: true, solution: false -->
```python
# your code should pass this test
assert bridge_formula(51, 5) == 79875
print('Success')
```

Output:
```
Success
```

<!-- grade_id: dot-test2, points: 1, grade: true, solution: false -->
```python
# your code should pass this test
assert bridge_formula(20,3) == 51000
print('Success')
```

Output:
```
Success
```

<div class="alert alert-info">
The interstate highway system was first established in the wake of World War I as a defense measure against ground invasion by a foreign power<sup>[[Wikipedia](https://en.wikipedia.org/wiki/Interstate_Highway_System#History)]</sup>.
</div>

### 6. How to divide your code logically

#### 6.1 Blocks

A *block* in Python refers to a section of code that is sequentially executed as a unit.  Blocks are differentiated by indenting the lines of code all together.  (And Python is *very* strict about the *number* of spaces being consistent.)  For instance, consider this program to calculate the pressure of an ideal gas given the temperature and volume (of given number of molecules):

$$
P \propto \frac{T}{V} = \frac{RT}{V}
\text{.}
$$

```python
def pressure_IG(T, V):
    # T should be in deg C and V in cubic meters
    R = 8.314 # ideal gas constant, joules / deg C
    P = R * T / V
    return P

temperature = 100.0 # deg C
volume = 0.01 # cubic meters
pressure = pressure_IG(temperature, volume)

print('The pressure of', volume, 'cubic meters of gas at', temperature, 'deg C is', pressure, 'pascals')
```

Output:
```
The pressure of 0.01 cubic meters of gas at 100.0 deg C is 83140.0 pascals
```

This program has two blocks.  The first is the outermost layer, the lines of code in the first column which will be executed sequentially.  Line by line, here is what the program does.  (Note that throughout, `print` should be using parentheses.)

First, Python notices that you have created a function `pressure_IG` that accepts two *arguments*, or independent variables.

![](./img/block-1.png)

The code within that function, however, is skipped over until later, and Python proceeds within *the same block* to the next line, defining a variable `temperature`.  (The comment line is ignored.)

![](./img/block-2.png)

Again, Python proceeds to interpret the next line as defining a variable `volume`.

![](./img/block-3.png)

Next, Python attempts to define a variable `pressure`.  But to do so, it finds that it needs to look at the block of code referred to by the function `pressure_IG`.  So Python carries the values of `temperature` and `volume` over and places them in `T` and `V`,

![](./img/block-4.png)

Given `T` and `V` already, Python now enters the block of the function definition for `pressure_IG` and executes the code there.  First `R` is defined...

![](./img/block-5.png)

...and then the ideal gas pressure `P` is calculated.

![](./img/block-6.png)

`pressure_IG` `return`s `P`...

![](./img/block-7.png)

...which value is then used to define `pressure` in the main block and output the final result. 

![](./img/block-8.png)

Blocks are differentiated by their indentation.  Often, blocks end with a specific command, like `return`.  However, this is not imperative.  For instance, some functions simply carry out a task and then end:

```python
def warn(msg):
    print('WARNING:  ', msg)

warn('Keyboard not responding. Press any key to continue.')
```

Output:
```
WARNING:   Keyboard not responding. Press any key to continue.
```

#### 6.2 Scope

We hinted earlier at the slightly strange situations that can arise.  Consider the script:

```python
def sqr(x):
    return x ** 2

x = 5
y = 3
print(sqr(y))
```

Output:
```
9
```

Here we see *two* variables `x`—one defined inside of the function `sqr` and the other in the main block of code.  We will use the concept of *variable scope* to understand what each `x` means and where.

Basically, if we have a single block of code, then any reference to a variable or name (such as `x`) is interpreted by Python to mean the `x` within that block.

We can do strange things with blocks:  since a variable defined inside of a block "overrides" other variables with the same name defined outside of the block, the following code is valid in Python, if confusing:

    x = 5

    def do_calc():
        x = 4           # <-- this is a different `x`!
        print(x)
    
    print(x)            # <-- this refers to the first `x` (with value 5)

<h4 style="color:#FF8C00">Exercises</h4>

What is the value of each `x` at the marked places in the code below?  (*I.e.*, after the line indicated has been executed.)

    x = 4
    x *= 2          # 1.
    
    def do_calc(x):
        print(x)    # 2.
        return x ** 2
    
    y = x + 2
    x = do_calc(y)  # 3.

<!-- grade_id: scope1, points: 1, grade: true, solution: true -->
```python
# answer 1 here; either "x = ###" or "###" is acceptable
x = 6
```

<!-- grade_id: scope2, points: 1, grade: true, solution: true -->
```python
# answer 2 here; either "x = ###" or "###" is acceptable
x = 8
```

<!-- grade_id: scope3, points: 1, grade: true, solution: true -->
```python
# answer 3 here; either "x = ###" or "###" is acceptable
x = 64
```

Again, this freedom means that you need to choose variable names responsibly:  `n_entries` is always better than `n`!

### 7. How to manipulate strings

One of the elegant notions of programming is that programming helps generalize the idea of function beyond just manipulating numbers.  What about other data types—in particular, the other data type we've examined thus far, the string?  Python does have a lot of basic string manipulation functions.

```python
myStr = 'Mr. Jock, TV quiz PhD, bags few lynx.  '
```

```python
len(myStr)              # counts number of characters in myStr
```

Output:
```
39
```

```python
myStr.count('x')        # counts number of occurrences of 'x' in myStr
```

Output:
```
1
```

```python
myStr.find('x')         # returns position of character 'x'
```

Output:
```
35
```

```python
myStr.lower()           # returns myStr as lower-case letters
```

Output:
```
'mr. jock, tv quiz phd, bags few lynx.  '
```

```python
myStr.upper()           # returns myStr as upper-case letters
```

Output:
```
'MR. JOCK, TV QUIZ PHD, BAGS FEW LYNX.  '
```

```python
myStr.title()           # returns myStr as title-case letters
```

Output:
```
'Mr. Jock, Tv Quiz Phd, Bags Few Lynx.  '
```

```python
myStr.replace('a', 'b') # replaces all occurrences of a with b in myStr
```

Output:
```
'Mr. Jock, TV quiz PhD, bbgs few lynx.  '
```

```python
myStr.strip()           # removes leading/trailing white space from myStr
```

Output:
```
'Mr. Jock, TV quiz PhD, bags few lynx.'
```

Note that `myStr` itself never changes—functions like `upper` *return a copy* of `myStr` with the requested change made.

String methods are also *composable*—that is, you can put them together to carry out more complex operations:

```python
myStr.strip().replace(' ', '-')
```

Output:
```
'Mr.-Jock,-TV-quiz-PhD,-bags-few-lynx.'
```

```python
myStr.replace(' ', '-').strip()
```

Output:
```
'Mr.-Jock,-TV-quiz-PhD,-bags-few-lynx.--'
```

<h4 style="color:#FF8C00">Exercises</h4>

Now try these out:

```python
quip = '"Sometimes I wonder whether the world is being run by smart people who are putting us on, or by imbeciles who really mean it." (Mark Twain)'
```

-   How long is `quip`?

<!-- grade_id: quip1, points: 1, grade: true, solution: true -->
```python
# your answer here
ans = len(quip)
print(ans)
```

Output:
```
139
```

-   How many times does `'b'` occur in this phrase?

<!-- grade_id: quip2, points: 1, grade: true, solution: true -->
```python
# your answer here
ans = quip.count('b')
print(ans)
```

Output:
```
4
```

-   How many times does `' '` occur?

<!-- grade_id: quip3, points: 1, grade: true, solution: true -->
```python
# your answer here
ans = quip.count(' ')
print(ans)
```

Output:
```
25
```

-   Replace all instances of `' '` with `'a'`.  Now how many times does `'a'` occur?  (You can do this using an intermediate variable or by composing the operations as above.)

<!-- grade_id: quip4, points: 1, grade: true, solution: true -->
```python
# your answer here
ans1 = quip.replace(' ', 'a')
print(ans1)
ans2 = ans1.count('a')
print(ans2)
```

Output:
```
"SometimesaIawonderawhetheratheaworldaisabeingarunabyasmartapeopleawhoaareaputtingausaon,aorabyaimbecilesawhoareallyameanait."a(MarkaTwain)
31
```

Finally, let's close out this section with a teaser for the `list`, which we still haven't discussed in detail yet.  Try these:

```python
quip.split(' ')
```

Output:
```
['"Sometimes',
 'I',
 'wonder',
 'whether',
 'the',
 'world',
 'is',
 'being',
 'run',
 'by',
 'smart',
 'people',
 'who',
 'are',
 'putting',
 'us',
 'on,',
 'or',
 'by',
 'imbeciles',
 'who',
 'really',
 'mean',
 'it."',
 '(Mark',
 'Twain)']
```

```python
myStr.split(',')
```

Output:
```
['Mr. Jock', ' TV quiz PhD', ' bags few lynx.  ']
```

These break the string along the fault lines specified, or the *delimiters* `' '` and `','`, and return the result as a collection of strings known as a list.

-   Convert the string `' Time present and time past / Are both perhaps present in time future '` into the string `'Tim3_Pr3S3Nt_And_Tim3_Past_/_Ar3_Both_P3Rhaps_Pr3S3Nt_In_Tim3_Futur3'` using only the above functions.  You should only require four functions, but may take more if necessary.  The final result should be stored in the variable `ans`.

<!-- grade_id: xform, points: 0, grade: false, solution: true -->
```python
# your code here
start = ' Time present and time past / Are both perhaps present in time future '
ans = start.strip().replace('s','S').replace('e','3').replace(' ', '_').title()
print(ans)
```

Output:
```
Tim3_Pr3S3Nt_And_Tim3_Past_/_Ar3_Both_P3Rhaps_Pr3S3Nt_In_Tim3_Futur3
```

<!-- grade_id: xform-test, points: 1, grade: true, solution: false -->
```python
# this is a test which your code should pass
assert ans == 'Tim3_Pr3S3Nt_And_Tim3_Past_/_Ar3_Both_P3Rhaps_Pr3S3Nt_In_Tim3_Futur3'
print('Success')
```

Output:
```
Success
```

### 8. Submission
1.  **File name**: save and rename file as `lab03_[your_student_ID].ipynb` after you're done and send it to `cs101homework@intl.zju.edu.cn`, **Must send files by your INTL email**.
2. **Email title**: Please use the following format: `CS101-Lab03-Group[A1]-[your_student_ID]`. Unclear email titles may result in a credit penalty.
3.  **Hand in the handout to the instructor at the end of this lab session**. Remember to write down your name, your student ID, and the lab section (lab#03) on the second page.
4. **You have 1 hour and 50 minutes to complete the lab and late submissions will not be tolerated**, since the labs are running at the capacity of the room.

```python

```