## `lab02`—Expressions & Data

**Objectives**

-   Understand various ways of writing text as strings.
-   Compose compound expressions to calculate mathematical formulae.

### How to write text

Numbers are never enough—you use them to get the math right, but when it comes time to communicating the results to the outside world, you need text.  In programming terms, we refer to a block of text as a *string* (from a string of characters).  Some messages are hard-coded, like error messages; others will be generated as necessary by the program.

Since we communicate with the Python intepreter using text commands, we also need a way to distinguish between instructions to the machine (code) and messages written as text (strings).  In other words, we need Python to distinguish between

```python
print(hello world)
```

and

```python
print('hello world')
```

The string holds the words within quotes as a single unit.  We typically write strings by surrounding our text in single quotes `''` or double quotes `""`:

```python
'All that glitters is not gold.'
```

```python
'Who knows what evil lurks in the hearts of men?  The Shadow knows.'
```

```python
'I'm sorry, Dave, I'm afraid I can't do that.'
```

That last one has a problem, doesn't it? We wrote a string by marking the ends with single quotes ', but then used a single quote inside of the string.  This thoroughly confused Python, as you can see by the red and black text mixed.

This is why Python lets you use either single or double quotes:

```python
"I'm afraid I can't do that."
```

```python
'They Call the Wind "Maria"'
```

Write the following phrases as Python strings.  (You can copy and paste the text itself into the string.)

- George Orwell's novel 1984

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
```python

```

- This is a straight double quote, ", while these are curly double quotes: “, ”.

<!-- grade_id: string2, points: 1, grade: true, solution: true -->
```python

```

If things get really hairy, you can also use triple-quotes to mark a string:

```python
'''She then turned and asked the librarian, "Do you have a copy of Diderot's and d'Alembert's Encyclopedie?"'''
```

<div class="alert alert-info">
Notice that the string returned by Python consists of *single quotes* wrapping a sentence containing *double quotes*, which then contain more *single quotes*.  The inner single quotes have a backslash in front of them—this is called an *escape character*, which is a fancy way of saying that the backslash tells Python to treat the single quote as a 'quote' rather than as marking the beginning or end of a string.
</div>

### How to store a value for future use

In physics and engineering, we often worked with labeled unknowns, called *variables*.  This lets us refer to physical or mathematical values without knowing the value ahead of time.  For instance, the ideal gas law represents the relationship between pressure $P$, temperature $T$, and volume $V$, written as

$$
PV \propto T \text{.}
$$

We can express the relationship between these physical quantities without referring to specific values.

For a computer, variables work much the same way—you can use them to write expressions and formulas, and you can use them to store the results of past calculations.  You create a variable using the assignment operator `=` and reuse it using the same name in subsequent expressions:

```python
x = 5
```

```python
x
```

```python
x**2
```

```python
x%2
```

```python
x**x
```

```python
x/x
```

You update the value of a variable by simply overwriting it:

```python
x = 4.5
```

```python
x
```

```python
x**2
```

In fact, you can even use a variable's *former* value to give it a *new* value:

```python
x = 2*x
```

If you missed it in the second lecture, stop and think about this for a moment, because it's very important.  If you were to write down the mathematical statement

$$
x = 2x \text{,}
$$

there is *no solution*—$x$ divides out of both sides and you are left with the mathematical absurdity $1=2$.  You should *not think* of the assignment operator as reading "equals".  That is, `x = 2x` should be read by you as "`x` is made equal to the value of 2 times `x`".

<div class="alert alert-warning">
The *assignment* operator is a single equals sign, `=`, read "is made equal to".
<br/>
The *equality* operator is a double equals sign, `==`, read "equals".
</div>

As you saw in the lecture, you can use much more involved variable names, which is a great idea if you don't want to lose track of what you're doing.

```python
trees_observed = 4
max_torque = 5.0
middleInitial = 'L'
finalSum = 107
way_to_Tipperary = 'east-north-east'
warning_msg = "I am Jack's medulla oblongata."
```

A few possible names do not, in fact, work:

```python
1st = 16.5
```

```python
$ = 2
```

<div class="alert alert-warning">
The error that occurs in both of these cases, `SyntaxError`, generally means that you've mistyped something or tried to express something impossible by the rules of Pythonic syntax. Basically, as long as you use only letters, numbers, and the underscore `_` in your variable names, you'll be fine.
</div>

### How variables make your life easier

Imagine writing a long script which depends on the value of a heat flux operator with the value $7.9 \,\text{watts}\cdot\text{meter}^{-2}\cdot\text{kelvin}$.  You dutifully write the statements to calculate the heat transfer through the wall with a difference of $\Delta T = 1.2 \,\text{kelvin}$ using the value `7.9`:

```python
print("The heat transfer is ", 7.9 * 1.2, " W/m**2.")
```

You may have already noticed the problem here—what happens when the problem changes slightly and you need to use a different heat flux operator value?  This is where variables shine:

$$
q = h \Delta T
$$

If you set things up properly, you will only have to define a value for $h$ and $\Delta T$ once, at the beginning of your script, and it will be the same throughout the program.  Otherwise, you might have had to search and replace every single instance—and hope you didn't miss one!

```python
h = 7.9     # heat transfer coef, W/m^2/deg C
DT = 1.2    # difference in temperature between fluid and surface, deg C
q = h * DT  # heat flux or amount of heat transferred, W/m^2

print("The heat transfer is ", q, " W/m**2.")

DT = 1.3
q = h * DT  # heat flux or amount of heat transferred, W/m^2

print("The heat transfer is ", q, " W/m**2.")
```

<div class="alert alert-info">
Note as well that since physical quantities always have physical dimensions, you should always document these with the value (or include a note to the effect that all values are [SI](https://en.wikipedia.org/wiki/International_System_of_Units), [CGS](https://en.wikipedia.org/wiki/Centimetre%E2%80%93gram%E2%80%93second_system_of_units), [FPS](https://en.wikipedia.org/wiki/Foot%E2%80%93pound%E2%80%93second_system), etc.).
<pre><code>max_torque = 5.0  # N*m</code></pre>
A brief explanation may be helpful as well:
<pre><code>dx = 1e-2  # mesh discretization, m</code></pre>
Form this habit now and it will serve you well as you begin writing code, scripts, and longer programs.
</div>

-   Correct the following code snippet to use variables `v` and `a` instead of hard-coded values.

<!-- grade_id: parameterize, points: 0, grade: false, solution: true -->
```python
# modify this code
# the base equation is:
# d = v*t + 0.5*a*t*t
t = 10.0
d = 5.0 * t + 0.5 * -9.8 * t ** 2
```

<!-- grade_id: parameterize-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
t = 10.0
d = v*t + 0.5*a*t**2
assert isclose(d, -440)
print('Success!')
```

### Building expressions using operators

You've seen basic expressions in the lab and in the homework.  Let's build some more complicated expressions now to serve us in real model-building later.

-   Write an expression describing the quadratic formula,
    $$
    x = -b + \frac{\sqrt{b^2-4ac}}{2a} \text{.}
    $$
    
    Use the following variable values in your solution:

```python
a = 1
b = 2
c = 1
```

<!-- grade_id: quad, points: 0, grade: false, solution: true -->
```python
# write a formula for x here
```

<!-- grade_id: quad-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
assert x == -2.0
print('Success!')
```

-   Write an expression describing the conversion of a degree/minute/second latitude value $d°m′s″$ into an equivalent decimal latitude value $\ell$,
    $$
    \ell = d + \frac{m}{60} + \frac{s}{3600} \text{.}
    $$
    
    Call your Python variable containing the result `lat` instead of `l`.

```python
d = 22
m = 30
s = 120
```

<!-- grade_id: dms, points: 0, grade: false, solution: true -->
```python
# write a formula for lat here
```

<!-- grade_id: dms-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(lat, 22.53333)
print('Success!')
```

-   Write an expression describing the period of a pendulum for small amplitudes,
    $$
    T = 2\pi \sqrt\frac{L}{g} \text{.}
    $$

```python
from math import pi
L = 0.72  # cm
g = 9.8   # acceleration due to gravity.
```

<!-- grade_id: pend, points: 0, grade: false, solution: true -->
```python
# write a formula for T here, using pi (and not an approximation)
```

<!-- grade_id: pend-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(T, 1.70307)
print('Success!')
```

A very common pattern is to need to add a value to a variable, and then to store the result of that addition as a new variable. Consider:
    
    sum = 0
    sum = sum + 5  # the value of sum is now 5
    sum = sum + 4  # the value of sum is now 5 + 3 = 8
    sum = sum + 3  # the value of sum is now 8 + 8 = 16

Since this type of statement occurs a lot, Python has a shorthand way of writing the statement:
    
    sum += 5

-   Using the `+=` and `-=` operators, add the following sequence of charges to the `charge` variable:
    $$
    +1, -2, -1, +2, +3
    $$

```python
charge = +2  # ion charge
```

<!-- grade_id: charge, points: 1, grade: false, solution: true -->
```python
# write your series of +=, -= statements here
```

<!-- grade_id: charge-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
assert charge == +5
print('Success!')
```