## `lab04`—Logic & Control

**Objectives**

-   Use functions effectively (with arguments, comments, and scripting).
-   Understand Python's binary logic system.
-   Exploit conditional statements to build adaptive programs.

Let's do some more work with functions to get used to their applications.

Imagine performing more complicated tasks like permuting a set of variables (which is useful in tensor algebra, for instance<sup>[[MathWorld](http://mathworld.wolfram.com/Tensor.html)]</sup>):

<img src="./img/permute.png" width="50%;"/>

-   Write a function `permute` which carries out this operation.  (It really doesn't need to be more complicated than the `swap` function above, just with more variables.)

<!-- grade_id: permute, points: 0, grade: false, solution: true -->
```python
# define your function here
def permute('''(delete this string and replace it with the incoming variables)'''):
    pass # you can always delete a `pass` statement, since it does nothing
```

<div class="alert alert-warning">
The code stub above has a `pass` statement.  This is a placeholder which does nothing, and you should replace it with either a `return` statement or another block of appropriate code.  Notebooks that you submit should never leave `pass` statements in.
</div>

<!-- grade_id: permute-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
i,j,k,l,m = 1,2,4,8,16
assert (m,i,j,k,l) == permute(i,j,k,l,m):
print('Success!')
```

### Optional Parameters

Occasionally, it will be convenient for us to assign default values to parameters.  Consider a function which calculates to some level of precision, or *tolerance*.  While it is good for the user to be able to specify when a very high tolerance is necessary (like $10^{-6}$), for most purposes it may be simpler to just have a default tolerance built-in.

This function determines if two numbers are within a certain range of each other, *i.e.*,

$$
\frac{|a-b|}{\max \left( |a|, |b| \right)} \leq \varepsilon_\textrm{tol}
$$

```python
def isclose(a, b, tol=1e-3):
    result = abs(a - b) <= tol * max(abs(a), abs(b))
    return result
```

In this case, `tol` defaults to `1e-3` *unless* the user specifies `tol`.  The following are all valid calls to `evaluate`:

```python
value1 = 5.000
value2 = 5.001

a = isclose(value1, value2)
b = isclose(value1, value2, tol=1e-4)
c = isclose(value1, value2, tol=1e-5)

print(a, 'for a tolerance of 1e-3 (default)')
print(b, 'for a tolerance of 1e-4')
print(c, 'for a tolerance of 1e-5')
```

-   Write a function `inv_power` which accepts two parameters, `x` and `b`.  `b` should supply a default value of `1.0` if it is not given as an argument.  The function should return the result $x^{-b}$.

<!-- grade_id: inv_power, points: 0, grade: false, solution: true -->
```python
# define your function here
def inv_power('''(delete this string and replace it with the incoming variables)'''):
    pass # you can always delete a `pass` statement, since it does nothing
```

<!-- grade_id: inv_power-spec, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for specified case
assert inv_power(3.0,3) == 0.037037037037037035
print('Success!')
```

<!-- grade_id: inv_power-default, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for default case
assert inv_power(3.0) == 0.3333333333333333
print('Success!')
```

<div class="alert alert-info">
This arrangement is very common in libraries like [Scientific Python (`scipy`)](https://scipy.org/scipylib/index.html) in which functions can have many parameters.  Consider the definition of the function `quad` which integrates a function:
<pre><code>scipy.integrate.quad(func,              # function to integrate
                     a,                 # left-hand limit
                     b,                 # right-hand limit
                     args=(),           # necessary function arguments
                     full_output=0,     # various parameters for precision and numerics
                     epsabs=1.49e-08,   # ...
                     epsrel=1.49e-08,
                     limit=50,
                     points=None,
                     weight=None,
                     wvar=None,
                     wopts=None,
                     maxp1=50,
                     limlst=50)
</code></pre>
<div>
The first three parameters `func`, `a`, and `b`, are required (have no defaults); all of the others have default values but can be overridden if the user wishes to do something specific with the calculation.
</div>
</div>

### Binary Logic & Comparison Operators

<div class="alert alert-info">
You saw this in the lecture last week, but this is a quick refresher.
</div>

Mathematically, when you wish to express whether or not two things are greater than, less than, or equal to each other, you write statements along the lines of
$$a > 5; b < 3; x = \frac{\pi}{3}; b \leq 10; a \neq b \text{.}$$

In the Python language, things are much the same.  The previous statements could be expressed thus:

| $\text{Mathematics}$ | `Python`  |
|----------------------|-----------|
| $a > 5$              | `a > 5`   |
| $b < 3$              | `b < 3`   |
| $x = \frac{\pi}{3}$  | `x == pi/3` |
| $b \leq 10$          | `b <= 10` |
| $a \neq b$           | `a != b`  |

If you define the variables in our example, you can try them out.  Python answers tersely with `True` or `False`.  These are know as the *Boolean* or *binary* types, and arise as the result of comparisons.

Comparisons can be *chained* together using `and` and `or` to create a single compound expression.  We won't delve deeply into the mysteries of binary logic, but this construction is useful when several conditions must be met.  To enforce *both* conditions of the first two lines of the above table, write

    a > 5 and b < 3

To enforce *either* one or *both*, write

    a > 5 or b < 3

-   Write a function `compare` which accepts three variables, `r`, `x`, and `s`, and returns the result of $r < x \leq s$ by a single compound expression.

<!-- grade_id: compare, points: 0, grade: false, solution: true -->
```python
# define your function here
def compare('''(delete this string and replace it with the incoming variables)'''):
    pass # you can always delete a `pass` statement, since it does nothing
```

<!-- grade_id: compare-spec1, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for specified case
r = 7.64e-3
x = 7.64e-3
s = 45.0
assert compare(r, x, s) == False
print('Success!')
```

<!-- grade_id: compare-spec2, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for specified case
r = 7.64e-3
x = 45.0
s = 45.0
assert compare(r, x, s) == True
print('Success!')
```

###  Conditional Statements & Branched Control

Conditional logic lets you deal with both *branches* of control.  There are two possible routes the program could take, like a river that splits around an island and then flows together again afterward.

Take, for instance, the definition of the absolute value:

$$
\text{abs}(x) =
|x| =
\left\{
\begin{array}{lr}
-x & : x < 0 \\
x & : x \geq 0
\end{array}
\right.
$$

Depending on the value of $x$, a different definition applies when evaluating the function $\text{abs}$.

```python
def abs(x):
    if x < 0:
        return -x
    else:
        return x

print('|-5| =', abs(-5.))
print('|+5| =', abs( 5.))
```

Last week you wrote a $\text{sinc}$ calculation function that included a single statement to deal with an edge case where a division by zero could occur.

$$\text{sinc}(x) = \frac{\sin(x)}{x}$$

With conditional statements and branch control (an `if` statement), it becomes trivial to deal with the edge case.

```python
def sinc(x):
    if x == 0:
        return 1.0  # safeguard against ZeroDivisionError
    y = sin(x) / x
    return y

sinc(0)
```

Essentially, the calculation of $\text{sinc}$ at any point falls into one of two categories:  either the calculation has a simple division by a real number or the taking of a limit is involved.  Since we already know the answer in all cases, we forgo the limit and simply `return 1.0` in this case.

<div class="alert alert-info">
Does Python treat comparisons with integers and floating-point numbers differently?  Try comparing a variable to `0` and to `0.0` below.
</div>

```python
# test if 0 and 0.0 behave similarly in comparisons
x = 0
if x == 0:
    print('Equal!')
```

### Application:  Polygon Areas

<img src="./img/apothem.png" width="25%;"/>

The area $A$ of a regular polygon is given by the formula

$$A = \frac{1}{2} a \times n \times s$$

where $a$ is the apothem, or the radius of an inscribed circle, $n$ is the number of sides, and $s$ is the length of a side.  A geometric relationship between the $a$, $s$, and $n$ is

$$a = \frac{s}{2 \tan(\pi / n)}$$

Together, we can use these facts to write an area calculation function with two inputs:  the number of sides and the length of a side.  A function `tan` and a constant `pi` are available in the module `math`.

<div class="alert alert-warning">
Python distinguishes between upper-case and lower-case letters—it is *case-sensitive*.  This means that `A` and `a` are *different variables*, just as we have used them in the previous mathematical expressions.
</div>

-   Compose a function `area_polygon` which calculates the area `A` of a regular polygon given the number of sides `n` and the length of a side `s` (with a default value of `1.0` for `s`).  (It should `return 0.0` if `n <= 2`.)  You should include any necessary `import` statements.

<!-- grade_id: area, points: 0, grade: false, solution: true -->
```python
# define your function here
def area_polygon('''(delete this string and replace it with the incoming variables)'''):
    pass # you can always delete a `pass` statement, since it does nothing
```

<!-- grade_id: area-default, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test default case
from numpy import isclose
assert isclose(area_polygon(7), 3.6339124440015893)
print('Success!')
```

<!-- grade_id: area-spec, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test specified case
from numpy import isclose
assert isclose(area_polygon(9, s=2.0), 24.727296775091602)
print('Success!')
```

<!-- grade_id: area-degen, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test degenerate case
from numpy import isclose
assert isclose(area_polygon(2), 0.0)
print('Success!')
```

<div class="alert alert-info">
The use of the apothem derives from ancient geometry, which was built on the <a href="https://en.wikipedia.org/wiki/Chord_(geometry)#crd">*chord*</a> of the circle rather than on the sine and cosine relations.
</div>

### Application:  Heat Transfer Correlation

A random opening of a textbook on heat and mass transfer suggests as an example the piecewise function describing convective heat transfer away from a hot plate<sup>[[Incropera2002, p. 557](#Incropera2002)]</sup>:

$$
h = 
\left\{
\begin{array}{lr}
0.54 Lk Ra_{\text{L}}^{1/4} & 10^4 \leq Ra_{\text{L}} \leq 10^7 \\
0.15 Lk Ra_{\text{L}}^{1/3} & 10^7 \leq Ra_{\text{L}} \leq 10^{11}
\end{array}
\right.
$$

where $h$ is the convective heat transfer coefficient, $L$ is the length of the plate, etc.  (Heat transfer is full of a lot of wonky correlations and compound definitions like this.)

You could write each of the right-hand-side comparisons as follows:

| $\text{Mathematics}$                | `Python`                     |
|-------------------------------------|------------------------------|
| $10^4 \leq Ra_{\text{L}} \leq 10^7$ | `1e4 <= Ra and Ra <= 1e7`  |
| $10^7 < Ra_{\text{L}} \leq 10^{11}$ | `1e7 < Ra and Ra <= 1e11` |

-   Write a function `heat_transfer` which accepts a number `Ra` and returns the heat transfer coefficient `h` as calculated from the above formula.  You should test whether the input values are between `1e4` and `1e11`; if they are outside this range, `heat_transfer` should `return None`.  Use a comparison (an `if` statement) to test which expression to use.

<!-- grade_id: heat, points: 0, grade: false, solution: true -->
```python
# define your function here
def heat_transfer('''(delete this string and replace it with the incoming variables)'''):
    L = 1.0
    k = 1.0
    pass # you can always delete a `pass` statement, since it does nothing
```

<!-- grade_id: heat-range, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for value out of range
assert heat_transfer(1e1) == None
print('Success!')
```

<!-- grade_id: heat-bottom, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for specified case
from numpy import isclose
assert isclose(heat_transfer(1e5), 9.602708814210184)
print('Success!')
```

<!-- grade_id: heat-top, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for specified case
from numpy import isclose
assert isclose(heat_transfer(1e8), 69.62383250419165)
print('Success!')
```