## `lab04`—Lists, Logic & Control

**Objectives**
-   Get familiar with List
-   Use functions effectively (with arguments, comments, and scripting).
-   Understand Python's binary logic system.
-   Exploit conditional statements to build adaptive programs.

### Lists

Many of the expressions you have written deal with single numbers or variables representing single values.  A notable exception is the string, which contains many single characters collected together.  The string is the first *compound data type* you have encountered.  Let's look at another:  the `list`.

In Python, `list`s are a general way of storing a collection of data together in one place for convenience.  Much like strings require quotes at the beginning and end, `list`s are marked by square brackets `[]`.  You can make `list`s containing any data type:

```python
[8,6,7,5309]
```

Output:
```
[8, 6, 7, 5309]
```

```python
[3.1415, 2.7183, 1.6180, 2.2195]
```

Output:
```
[3.1415, 2.7183, 1.618, 2.2195]
```

```python
['a', 'bcd', 'efghij']
```

Output:
```
['a', 'bcd', 'efghij']
```

Notably, you can mix data types:

```python
[1, 'abracadabra', None, abs]
```

Output:
```
[1, 'abracadabra', None, <function abs>]
```

Other expression, like `range`, can easily be turned into `list`s:

```python
range(10)
```

Output:
```
range(0, 10)
```

```python
list(range(10))
```

Output:
```
[0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
```

```python
list('asdf')
```

Output:
```
['a', 's', 'd', 'f']
```

This is a good time to remind you that if you need the entire range (including the right-hand limit), you need to add 1.  (This is commonly the case with grids for numerical problems, for instance.)

```python
nx = 15
list(range(nx))
```

Output:
```
[0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14]
```

```python
list(range(nx+1))
```

Output:
```
[0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15]
```

List indexing works much the same as string indexing, except that each element of the list is referred to rather than each digit or character.

```python
data = list(range(-10,10))
print(data)
```

Output:
```
[-10, -9, -8, -7, -6, -5, -4, -3, -2, -1, 0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
```

```python
data[1:-1]
```

Output:
```
[-9, -8, -7, -6, -5, -4, -3, -2, -1, 0, 1, 2, 3, 4, 5, 6, 7, 8]
```

```python
data[:-1]
```

Output:
```
[-10, -9, -8, -7, -6, -5, -4, -3, -2, -1, 0, 1, 2, 3, 4, 5, 6, 7, 8]
```

A neat trick—to reverse a list, index it from beginning to end and then count backwards:

```python
data[::-1]
```

Output:
```
[9, 8, 7, 6, 5, 4, 3, 2, 1, 0, -1, -2, -3, -4, -5, -6, -7, -8, -9, -10]
```

But when you run the method `reverse`, nothing seems to happen.

```python
data.reverse()
```

This is because `reverse` *changes the string itself*, rather than returning a copy of the contents with changes made.

```python
print(data)
```

Output:
```
[9, 8, 7, 6, 5, 4, 3, 2, 1, 0, -1, -2, -3, -4, -5, -6, -7, -8, -9, -10]
```

What happens when you multiply a `list`?

```python
print([5,10] * 3)
```

Output:
```
[5, 10, 5, 10, 5, 10]
```

<div class="alert alert-info">
If this dissatisfies you, and you really want the result of this last operation to be `[15, 30]`, then wait a few weeks until we introduce the [Numerical Python (`numpy`) module](https://docs.scipy.org/doc/numpy/user/whatisnumpy.html).  That will treat *arrays* as collections of numerical data, much like columns in a spreadsheet.
</div>

You can add (`append`) elements to or `remove` elements from a list:

```python
my_list = ['sunshine', 'breeze', 'daisy', 'clover', 'gasoline']
my_list.remove('gasoline')  # one of these things is not like the others!
print(my_list)
```

Output:
```
['sunshine', 'breeze', 'daisy', 'clover']
```

```python
my_list.append('napping')
print(my_list)
```

Output:
```
['sunshine', 'breeze', 'daisy', 'clover', 'napping']
```

### Functions
Now let's do some more work with functions to get used to their applications.

#### Optional Parameters

Occasionally, it will be convenient for us to assign default values to parameters.  Consider a function which calculates to some level of precision, or *tolerance*.  While it is good for the user to be able to specify when a very high tolerance is necessary (like $10^{-6}$), for most purposes it may be simpler to just have a default tolerance built-in.

This function determines if two numbers are within a certain range of each other, *i.e.*,

$$
\frac{|a-b|}{\max \left( |a|, |b| \right)} \leq \varepsilon_\textrm{tol}
$$

```python
def isclose(a, b, tol=1e-3):
    result = abs(a - b) <= tol * max(abs(a), abs(b))
    return result
```

In this case, `tol` defaults to `1e-3` *unless* the user specifies `tol`.  The following are all valid calls to `evaluate`:

```python
value1 = 5.000
value2 = 5.001

a = isclose(value1, value2)
b = isclose(value1, value2, tol=1e-4)
c = isclose(value1, value2, tol=1e-5)

print(a, 'for a tolerance of 1e-3 (default)')
print(b, 'for a tolerance of 1e-4')
print(c, 'for a tolerance of 1e-5')
```

Output:
```
True for a tolerance of 1e-3 (default)
False for a tolerance of 1e-4
False for a tolerance of 1e-5
```

-   Write a function `inv_power` which accepts two parameters, `x` and `b`.  `b` should supply a default value of `1.0` if it is not given as an argument.  The function should return the result $x^{-b}$.

<!-- grade_id: inv_power, points: 0, grade: false, solution: true -->
```python
# define your function here
def inv_power(x,b=1.0):
    return x**(-b)
    pass # you can always delete a `pass` statement, since it does nothing
```

<!-- grade_id: inv_power-spec, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for specified case
assert inv_power(3.0,3) == 0.037037037037037035
print('Success!')
```

Output:
```
Success!
```

<!-- grade_id: inv_power-default, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for default case
assert inv_power(3.0) == 0.3333333333333333
print('Success!')
```

Output:
```
Success!
```

<div class="alert alert-info">
This arrangement is very common in libraries like [Scientific Python (`scipy`)](https://scipy.org/scipylib/index.html) in which functions can have many parameters.  Consider the definition of the function `quad` which integrates a function:
<pre><code>scipy.integrate.quad(func,              # function to integrate
                     a,                 # left-hand limit
                     b,                 # right-hand limit
                     args=(),           # necessary function arguments
                     full_output=0,     # various parameters for precision and numerics
                     epsabs=1.49e-08,   # ...
                     epsrel=1.49e-08,
                     limit=50,
                     points=None,
                     weight=None,
                     wvar=None,
                     wopts=None,
                     maxp1=50,
                     limlst=50)
</code></pre>
<div>
The first three parameters `func`, `a`, and `b`, are required (have no defaults); all of the others have default values but can be overridden if the user wishes to do something specific with the calculation.
</div>
</div>

### Binary Logic & Comparison Operators

Mathematically, when you wish to express whether or not two things are greater than, less than, or equal to each other, you write statements along the lines of
$$a > 5; b < 3; x = \frac{\pi}{3}; b \leq 10; a \neq b \text{.}$$

In the Python language, things are much the same.  The previous statements could be expressed thus:

| $\text{Mathematics}$ | `Python`  |
|----------------------|-----------|
| $a > 5$              | `a > 5`   |
| $b < 3$              | `b < 3`   |
| $x = \frac{\pi}{3}$  | `x == pi/3` |
| $b \leq 10$          | `b <= 10` |
| $a \neq b$           | `a != b`  |

If you define the variables in our example, you can try them out.  Python answers tersely with `True` or `False`.  These are know as the *Boolean* or *binary* types, and arise as the result of comparisons.

Comparisons can be *chained* together using `and` and `or` to create a single compound expression.  We won't delve deeply into the mysteries of binary logic, but this construction is useful when several conditions must be met.  To enforce *both* conditions of the first two lines of the above table, write

    a > 5 and b < 3

To enforce *either* one or *both*, write

    a > 5 or b < 3

-   Write a function `compare` which accepts three variables, `r`, `x`, and `s`, and returns True if $r < x$ and $x \leq s$ (and False otherwise) using a compound expression.

<!-- grade_id: compare, points: 0, grade: false, solution: true -->
```python
# define your function here
def compare(r,x,s):
    return r < x and x <= s
    pass # you can always delete a `pass` statement, since it does nothing
```

<!-- grade_id: compare-spec1, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for specified case
r = 7.64e-3
x = 7.64e-3
s = 45.0
assert compare(r, x, s) == False
print('Success!')
```

Output:
```
Success!
```

<!-- grade_id: compare-spec2, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for specified case
r = 7.64e-3
x = 45.0
s = 45.0
assert compare(r, x, s) == True
print('Success!')
```

Output:
```
Success!
```

###  Conditional Statements & Branched Control

Conditional logic lets you deal with both *branches* of control.  There are two possible routes the program could take, like a river that splits around an island and then flows together again afterward.

Take, for instance, the definition of the absolute value:

$$
\text{abs}(x) =
|x| =
\left\{
\begin{array}{lr}
-x & : x < 0 \\
x & : x \geq 0
\end{array}
\right.
$$

Depending on the value of $x$, a different definition applies when evaluating the function $\text{abs}$.

```python
def abs(x):
    if x < 0:
        return -x
    else:
        return x

print('|-5| =', abs(-5.))
print('|+5| =', abs( 5.))
```

Output:
```
|-5| = 5.0
|+5| = 5.0
```

Last week you wrote a $\text{sinc}$ calculation function that included a single statement to deal with an edge case where a division by zero could occur.

$$\text{sinc}(x) = \frac{\sin(x)}{x}$$

With conditional statements and branch control (an `if` statement), it becomes trivial to deal with the edge case.
Now write the $\text{sinc}$ function that returns value 1 when $x = 0$, and does the normal calculation when $x \neq 0$

```python
from math import sin
def sinc(x):
    if x == 0:
        return 1
    return sin(x)/x
    pass # you can always delete a `pass` statement, since it does nothing
```

```python
assert sinc(0) == 1
print('Success!')

sinc(3.14)
```

Output:
```
Success!
0.0005072143046136395
```

Essentially, the calculation of $\text{sinc}$ at any point falls into one of two categories:  either the calculation has a simple division by a real number or the taking of a limit is involved.  Since we already know the answer in all cases, we forgo the limit and simply `return 1.0` in this case.

<div class="alert alert-info">
Does Python treat comparisons with integers and floating-point numbers differently?  Try comparing a variable to `0` and to `0.0` below.
</div>

```python
# test if 0 and 0.0 behave similarly in comparisons
x = 0.0
if x == 0:
    print('Equal!')
```

Output:
```
Equal!
```

### Application I:  Polygon Areas

<img src="./img/apothem.png" width="25%;"/>

The area $A$ of a regular polygon is given by the formula

$$A = \frac{1}{2} a \times n \times s$$

where $a$ is the apothem, or the radius of an inscribed circle, $n$ is the number of sides, and $s$ is the length of a side.  A geometric relationship between the $a$, $s$, and $n$ is

$$a = \frac{s}{2 \tan(\pi / n)}$$

Together, we can use these facts to write an area calculation function with two inputs:  the number of sides and the length of a side.  A function `tan` and a constant `pi` are available in the module `math`.

<div class="alert alert-warning">
Python distinguishes between upper-case and lower-case letters—it is *case-sensitive*.  This means that `A` and `a` are *different variables*, just as we have used them in the previous mathematical expressions.
</div>

-   Compose a function `area_polygon` which calculates the area `A` of a regular polygon given the number of sides `n` and the length of a side `s` (with a default value of `1.0` for `s`).  (It should `return 0.0` if `n <= 2`.)  You should include any necessary `import` statements.

<!-- grade_id: area, points: 0, grade: false, solution: true -->
```python
# define your function here
def area_polygon(n,s = 1.0):
    if n<=2:
        return 0.0
    from math import tan
    from math import pi
    return 0.25*n*s*s/tan(pi/n)
    pass # you can always delete a `pass` statement, since it does nothing
```

<!-- grade_id: area-default, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test default case
from numpy import isclose
assert isclose(area_polygon(7), 3.6339124440015893)
print('Success!')
```

Output:
```
Success!
```

<!-- grade_id: area-spec, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test specified case
from numpy import isclose
assert isclose(area_polygon(9, s=2.0), 24.727296775091602)
print('Success!')
```

Output:
```
Success!
```

<!-- grade_id: area-degen, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test degenerate case
from numpy import isclose
assert isclose(area_polygon(2), 0.0)
print('Success!')
```

Output:
```
Success!
```

<div class="alert alert-info">
The use of the apothem derives from ancient geometry, which was built on the <a href="https://en.wikipedia.org/wiki/Chord_(geometry)#crd">*chord*</a> of the circle rather than on the sine and cosine relations.
</div>

### Application II:  Heat Transfer Correlation

A random opening of a textbook on heat and mass transfer suggests as an example the piecewise function describing convective heat transfer away from a hot plate<sup>[[Incropera2002, p. 557](#Incropera2002)]</sup>:

$$
h = 
\left\{
\begin{array}{lr}
0.54 Lk Ra^{1/4} & 10^4 \leq Ra \leq 10^7 \\
0.15 Lk Ra^{1/3} & 10^7 \leq Ra \leq 10^{11}
\end{array}
\right.
$$

where $h$ is the convective heat transfer coefficient, $L$ is the length of the plate, etc.  (Heat transfer is full of a lot of wonky correlations and compound definitions like this.)

You could write each of the right-hand-side comparisons as follows:

| $\text{Mathematics}$                | `Python`                     |
|-------------------------------------|------------------------------|
| $10^4 \leq Ra \leq 10^7$ | `1e4 <= Ra and Ra <= 1e7`  |
| $10^7 < Ra \leq 10^{11}$ | `1e7 < Ra and Ra <= 1e11` |

-   Write a function `heat_transfer` which accepts a number `Ra` and returns the heat transfer coefficient `h` as calculated from the above formula.  You should test whether the input values are between `1e4` and `1e11`; if they are outside this range, `heat_transfer` should `return None`.  Use a comparison (an `if` statement) to test which expression to use.

<!-- grade_id: heat, points: 0, grade: false, solution: true -->
```python
# define your function here
def heat_transfer(Ra):
    L = 1.0
    k = 1.0
    if 1e4 <= Ra and Ra <= 1e7:
        return 0.54*L*k*Ra**(1/4)
    elif 1e7 < Ra and Ra <= 1e11:
        return 0.15*L*k*Ra**(1/3)
    else:
        return None
    pass # you can always delete a `pass` statement, since it does nothing
```

<!-- grade_id: heat-range, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for value out of range
assert heat_transfer(1e1) == None
print('Success!')
```

Output:
```
Success!
```

<!-- grade_id: heat-bottom, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for specified case
from numpy import isclose
assert isclose(heat_transfer(1e5), 9.602708814210184)
print('Success!')
```

Output:
```
Success!
```

```python
# it should pass this test---do NOT edit this cell
# test for specified case
from numpy import isclose
assert isclose(heat_transfer(1e8), 69.62383250419165)
print('Success!')
```

Output:
```
Success!
```

<!-- grade_id: heat-top, points: 1, grade: true, solution: false -->
```python
# It should print None ---do NOT edit this cell
# test for specified case
print(heat_transfer(1e2))
print(heat_transfer(1e13))
```

Output:
```
None
None
```