## `lab9`—Modeling Fluid Behavior

❖ Objectives

-   Use $n$-dimensional arrays to manage state data.
-   Plot data.
-   Create a model of scientific data.

In today's lab, you will implement various three fluid modeling laws.  You will then use these data to compare their results with real data and assess the relative accuracy of the various models.

##  Ideal Gas Law

In the physical description of liquids and gases, it often becomes necessary to use a relationship between the pressure on a volume of fluid and its temperature.  The simplest of these relationships is the *ideal gas law*,

$$PV = RT$$

(for a specified number of molecules of the gas).

The ideal gas law describes gases very well near atmospheric conditions and at lower pressures, but deviates from the true behavior of gases under other circumstances.  Since accurate descriptions of fluid flow, chemical reactivity, and oil field dynamics (among other things) depend on describing the state of a liquid, corrections are required.

1.  Write the ideal gas law as a function `pig` which accepts parameters `V` in cubic meters and `T` in degrees Kelvin (absolute temperature) and returns the corresponding `P`ressure in pascals.  $R = 8.314 \,\textrm{joule}\cdot\textrm{K}$, a proportionality constant.

<!-- grade_id: pig, points: 0, grade: false, solution: true -->
```python
# define your function here
def pig(V, T):
    """
    Return the ideal-gas pressure for a specified volume and temperature.
    
    Keyword arguments:
    V -- volume, in cubic meters
    T -- temperature, in Kelvin
    """
    
    ## YOU WRITE THIS BLOCK
    
    # Finally, return the ideal-gas pressure.
    return P
```

You may test your function against the following values:
    
| $P$/Pa | $V$/m$^3$ | $T$/K |
|--------|-----------|-------|
| 83140  | 0.01      | 100   |
| 166280 | 0.01      | 200   |

```python
# test your code here
pig(0.01, 273)
```

<!-- grade_id: pig-test1, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(pig(0.01, 100), 83140)
print('Success!')
```

<!-- grade_id: pig-test2, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(pig(0.01, 200), 166280)
print('Success!')
```

##  Cubic Equations of State

The true behavior of gases is extremely complex, but several equations have been introduced to model the gas characteristics approximately.  Among these so-called *equations of state* are the van der Waals equation of state<sup>[[Wikipedia](https://en.wikipedia.org/wiki/Equation_of_state#Van_der_Waals_equation_of_state)]</sup>,

$$
\begin{align*}
P_{\text{vdW}} & = \frac{RT}{V-b} - \frac{a}{V^2} \\
\end{align*}
$$

and the Redlich–Kwong equation of state<sup>[[PhysicsWiki](http://physwiki.ucdavis.edu/Condensed_Matter/Equations_of_state/Redlich-Kwong_equation_of_state)]</sup>,

$$
\begin{align*}
P_{\text{RK}} & = \frac{RT}{V-b} - \frac{a}{\sqrt{T}V(V+b)} \text{.} \\
\end{align*}
$$

(Constants $a$ and $b$ are specific to the gas being modeled:  that is, argon has a certain $a_\text{Ar}$, $b_\text{Ar}$; oxygen has a certain $a_{\text{O}_2}$, $b_{\text{O}_2}$, and so forth.  These typically use constants from the specific fluid known as the *critical temperature* $T_c$, *critical pressure* $P_c$, and *critical volume* $V_c$.

1.  Write the van der Waals equation of state as a function `pvdw` which accepts parameters `V`, `T`, `Tc`, and `Pc`, and returns the corresponding `P`ressure.
    
$$
P_{\text{vdW}} = \frac{RT}{V-b} - \frac{a}{V^2}
\hspace{2cm}
a = \frac{27}{64} \frac{\left(R T_c \right)^{2}}{P_c}
\hspace{2cm}
b = \frac{R T_c}{8 P_c}
\text{.}
$$

<!-- grade_id: pvdw, points: 0, grade: false, solution: true -->
```python
# define your function here
def pvdw('''add parameters here'''):
    """
    Return the van der Waals pressure for a specified volume and temperature.
    
    Keyword arguments:
    V -- volume, in cubic meters
    T -- temperature, in Kelvin
    Tc -- critical temperature, in Kelvin
    Pc -- critical pressure, in Kelvin
    """
    
    # Define parameters a and b.
    ## YOU WRITE THIS BLOCK
    
    # Calculate the vdw pressure.
    ## YOU WRITE THIS BLOCK
    
    # Finally, return the vdw pressure.
    return P
```

You may test your function against the following values:
    
| $P$ | $V$ | $T$ | $T_c$ | $P_c$ |
|-----|-----|-----|-------|-------|
| 40057.7 | 0.056634 | 273.15 | 190.6 | 4599000 |
| 47403.4 | 0.056634 | 323.15 | 190.6 | 4599000 |

```python
# test your code here
pvdw(0.01, 273, Tc=190.6, Pc=4599000)
```

<!-- grade_id: pvdw-test1, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(pvdw(0.056634, 273.15, Tc=190.6, Pc=4599000), 40057.7)
print('Success!')
```

<!-- grade_id: pvdw-test2, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(pvdw(0.056634, 323.15, Tc=190.6, Pc=4599000), 47403.4)
print('Success!')
```

1.  Write the Redlich–Kwong equation of state as a function `prk` which accepts parameters `V`, `T`, `Tc`, and `Pc`, and returns the corresponding `P`ressure.  You should `import` `sqrt` from `numpy` (not from `math`) or use the `**0.5` notation.

$$
P_{\text{RK}} = \frac{RT}{V-b} - \frac{a}{\sqrt{T}V(V+b)}
\hspace{2cm}
a = 0.42748 \frac{R^{2} T_c^{2.5}}{P_c}
\hspace{2cm}
b = 0.08662 \frac{R T_c}{P_c}
\text{.}
$$

<!-- grade_id: prk, points: 0, grade: false, solution: true -->
```python
# define your function here
def prk('''add parameters here'''):
    """
    Return the Redlich-Kwong pressure for a specified volume and temperature.
    
    Keyword arguments:
    V -- volume, in cubic meters
    T -- temperature, in Kelvin
    Tc -- critical temperature, in Kelvin
    Pc -- critical pressure, in Kelvin
    """
    
    # Define parameters a and b.  Note that a and b are different for prk and pvdw.
    ## YOU WRITE THIS BLOCK
    
    # Calculate the rk pressure.
    ## YOU WRITE THIS BLOCK
    
    # Finally, return the rk pressure.
    return P
```

You may test your function against the following values:
    
| $P$ | $V$ | $T$ | $T_c$ | $P_c$ |
|-----|-----|-----|-------|-------|
| 40059.4 | 0.056634 | 273.15 | 190.6 | 4599000 |
| 47408.3 | 0.056634 | 323.15 | 190.6 | 4599000 |

```python
# test your code here
prk(0.01, 273, Tc=190.6, Pc=4599000)
```

<!-- grade_id: prk-test1, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(prk(0.056634, 273.15, Tc=190.6, Pc=4599000), 40059.4)
print('Success!')
```

<!-- grade_id: prk-test2, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(prk(0.056634, 323.15, Tc=190.6, Pc=4599000), 47408.3)
print('Success!')
```

##  Modeling Data

As you can see, there are subtle variations between these equations of state even at fairly low temperature and pressure.  Let's see what these do for various fluids.  Since we are working with various fluids with different properties, let's set up a dictionary to store the fluid parameters, in particular the critical temperature and pressure.

Specifically, use a `dict` of `dict`s.  The outer `dict` has keys from the `Gas` column of the following table.  Each key has the value of another `dict`, this second `dict` containing keys `'Tc'` and `'Pc'` which contain the critical temperature and critical pressure, respectively.

|  Gas  | $T_c$ |  $P_c$  |
|-------|-------|---------|
| `H2`    | 33.19 | 1313000 |
| `O2`    | 154.6 | 5043000 |
| `N2`    | 126.2 | 3400000 |
| `CO2`   | 304.2 | 7383000 |
| `Ar`    | 150.9 | 4894000 |
| `C4H10` | 408.1 | 3648000 |

To make this a repeatable process, write it as a function.

-   Compose a function `gasdata` which accepts no parameters and `return`s a `dict` of `dict`s as described in the preceding paragraphs of this section.

<!-- grade_id: gasdata, points: 0, grade: false, solution: true -->
```python
# define your function here
def gasdata():
    """
    Return the van der Waals pressure for a specified volume and temperature.
    
    Keyword arguments:
    V -- volume, in cubic meters
    T -- temperature, in Kelvin
    Tc -- critical temperature, in Kelvin
    Pc -- critical pressure, in Kelvin
    """
    
    # Create the dictionary.
    gas_dict = {}
    gas_dict['H2'] = {'Tc': 33.19, 'Pc': 1313000}
    ## YOU WRITE THE REST OF THIS BLOCK
    
    # Finally, return the data dictionary.
    return gas_dict
```

```python
# test your code here
fluids = gasdata()
fluids['H2']  # try the others as well to test
```

<!-- grade_id: gasdata-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
try:
    fluid_test = gasdata()
    fluid_test['H2']['Tc']
    fluid_test['O2']['Pc']
    fluid_test['N2']['Tc']
    fluid_test['CO2']['Pc']
    fluid_test['Ar']['Tc']
except KeyError:
    raise AssertionError
print('Success!')
```

Now we can make our code more legible (always desirable):

```python
fluids = gasdata()

temp = 100 + 273.15  # 100 degC in K
vol  = 0.01  # m**3
pressure = prk(vol, temp, Tc=fluids['H2']['Tc'], Pc=fluids['H2']['Pc'])  # Pa
print(pressure)
```

Armed with

-   several models of gas behavior (the equations of state)
-   a database of fluid parameters

we can make and test predictions.

For ease-of-use, we will use NumPy arrays to store our data values.  (This allows us to perform *array operations*, that is, to multiply all elements in an array by a single value at the same time.)

Since the NumPy lecture will happen during this week, let's take a quick look at how arrays work in Python.

```python
import numpy as np  # import the numpy library; for convenience, we rename it as something a bit easier to type
np.set_printoptions(precision=5)  # sets floating-point number display to use only 5 digits after decimal point
```

A NumPy array is a collection of numeric data that acts together as a unit.  (Think of a spreadsheet column, for instance.)

![](./img/ndarray.png)

The nice thing is that the array behaves much more intuitively than `list`s sometimes do:

```python
my_list = [1, 2, 3, 4]
print('list * 3:', my_list * 3)

my_array = np.array(my_list)
print('arrayt * 3:', my_array * 3)
```

Output:
```
list * 3: [1, 2, 3, 4, 1, 2, 3, 4, 1, 2, 3, 4]
arrayt * 3: [ 3  6  9 12]
```

So arrays are like mathematical vectors, or spreadsheet columns, or however you'd like to think of them.  The point being, they can be manipulated as a unit which allows us to calculate many values at the same time (rather than singly in a `for` loop).

Contrast:

```python
new_list = []
for value in my_list:
    new_list.append(value * 2.0)
print(new_list)
```

Output:
```
[2.0, 4.0, 6.0, 8.0]
```

with the much sleeker:

```python
new_array = my_array * 2.0
print(new_array)
```

Output:
```
[ 2.  4.  6.  8.]
```

NumPy has become ubiquitous in scientific computing with Python for good reason, and we'll use it a lot in modeling over the next few weeks.

##  Modeling Fluids

One potential reason to use equations of state is to assess the storage pressure in gas cylinders at various temperatures, thus allowing safety officials to set guidelines.  A representative cylinder can hold 44 liters, or $0.044\,\text{m}^3$ (use this latter value).

In order to test the pressure across a range of temperatures, you should create an array of temperature points

```python
temps = np.array((273.15,  298.15,  323.15,  348.15,  373.15,  398.15))
```

and then store the results for pressure for each data point in an appropriate array variable.  For instance, for the Redlich–Kwong equation of state, the results could be calculate as

```python
fluids = gasdata()

volume  = 0.044  # m**3
pcyl_rk = prk(volume, temps, fluids['O2']['Tc'], fluids['O2']['Pc'])
print(pcyl_rk)
```

In addition to viewing the data in a tabular format, it is often far more intuitive to depict the relationship graphically.

```python
# these are the modules and command you need to plot in Jupyter
import matplotlib as mpl
import matplotlib.pyplot as plt
%matplotlib inline
```

```python
plt.plot(temps, pcyl_rk)  # define a basic plot, with x-values (temp) and y-values (pcyl_rk)
plt.show()  # show the plot
```

It will be convenient to measure the percent difference between two values,
    
$$
\%_\text{error} = \left|\frac{s - t}{s}\right| \times 100
$$
    
-   Compose a function `error` which accepts two `float`s `s` and `t` and returns the error as a percentage difference.  If you need to `import` any functions, use the `abs`olute value from `numpy`.

<!-- grade_id: error, points: 0, grade: false, solution: true -->
```python
# define your function here
def error(s, t):
    """
    Return the percentage difference between two values.
    
    Keyword arguments:
    s, t -- numeric values to compare, with reference to the former.
    """
    
    ## YOU WRITE THE REST OF THIS BLOCK
    
    # Finally, return the percentage difference (the 'error').
    return err
```

<!-- grade_id: error-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(error(4,5), 25)
print('Success!')
```

At the right-most data point in the last data calculation (temperature $398.15\,\text{K}$, volume $0.044\,\text{m}^3$), what is the percent difference between the ideal-gas law and Redlich–Kwong calculations, using $P_\text{rk}$ as the standard?  (We are modeling oxygen `'O2'` right now.)

Directly:

```python
Pig = pig(0.044, 398.15)
Prk = prk(0.044, 398.15, Tc=fluids['O2']['Tc'], Pc=fluids['O2']['Pc'])
print(error(Prk, Pig))
```

Using the former array:

```python
fluids = gasdata()

volume  = 0.044  # m**3
temps = np.array((273.15,  298.15,  323.15,  348.15,  373.15,  398.15))

pcyl_rk = prk(volume, temps, fluids['O2']['Tc'], fluids['O2']['Pc'])
pcyl_ig = pig(volume, temps)
print(error(pcyl_rk[-1], pcyl_ig[-1]))
```

The behavior is straightforward, just like with a list.

The real power occurs when you want to know more than just a single point:

```python
print(error(pcyl_rk, pcyl_ig))
```

NumPy permits `error` to be applied over the entire range at one go.  You can plot it as well:

```python
errors = error(pcyl_rk, pcyl_ig)
plt.plot(temps, errors)  # define a basic plot, with x-values (temp) and y-values (pcyl_rk)
plt.show()  # show the plot
```

Using the tools we've developed thus far, let's examine our use-case scenario of gases stored in a standard cylinder.  To refresh:

```python
fluids = gasdata()

volume  = 0.044  # m**3
temps = np.array((273.15,  298.15,  323.15,  348.15,  373.15,  398.15))

pcyl_ig = pig(volume, temps)
pcyl_rk = prk(volume, temps, Tc=fluids['O2']['Tc'], Pc=fluids['O2']['Pc'])
```

You should also define the van der Waals pressure as `pcyl_vdw`.

```python
# YOU DEFINE pcyl_vdw HERE
```

-   Plot all three gas laws across the range of temperatures.  (You can modify the following code to do so.)

<!-- grade_id: plt-all1, points: 1, grade: true, solution: true -->
```python
plt.plot(temps, pcyl_ig,
         temps, pcyl_rk)
plt.show()
```

-   Drop the volume by 90% to `0.0044` and calculate pressures at the same temperatures as above.  (Use the same variable names for `pcyl_ig`, etc.)  Plot the resulting error curves of `pcyl_ig` v. `pcyl_rk` and `pcyl_vdw` v. `pcyl_rk` across the range of temperature.  (There will be two curves in your plot.)

<!-- grade_id: plt-error1, points: 1, grade: true, solution: true -->
```python
# YOU RECALCULATE THE VARIABLES HERE

# YOU WRITE THE plt.plot STATEMENT
plt.show()
```

Argon is considered to behave very closely to an ideal gas because it is a spherical one-atom molecule.

Consider the pressure of argon in a small gas cylinder:

```python
fluids = gasdata()

volume = 0.044 # m**3
temps = np.array((273.15,  298.15,  323.15,  348.15,  373.15,  398.15))

pcyl_ig  = pig( volume, temps)
pcyl_rk  = prk( volume, temps, Tc=fluids['Ar']['Tc'], Pc=fluids['Ar']['Pc'])
pcyl_vdw = pvdw(volume, temps, Tc=fluids['Ar']['Tc'], Pc=fluids['Ar']['Pc'])

plt.plot(temps, pcyl_ig,  'g-',
         temps, pcyl_rk,  'r-',
         temps, pcyl_vdw, 'b-')
plt.show()
```

-   Plot the resulting error curves of `pcyl_ig` v. `pcyl_rk` and `pcyl_vdw` v. `pcyl_rk` across the range of temperature.  (There will be two curves in your plot.)

<!-- grade_id: plt-argon1, points: 1, grade: true, solution: true -->
```python
# YOU WRITE THE plt.plot STATEMENT AND ANY NECESSARY CODE
plt.show()
```

Argon's behavior doesn't deviate much from ideality across this range of temperatures.

In contrast, butane ($\text{C}_4\text{H}_{10}$, `C4H10`) is shaped like a Y and tends to behave very non-ideally.  It is a gas at room temperature and standard atmospheric pressure, but even in a tiny thin-shelled plastic lighter it is a liquid.  While we won't quantify the liquid–gas transition directly (which is a pretty involved calculation), we will see how far from ideality butane is.

```python
fluids = gasdata()

volume = 0.044 # m**3
temps = np.array((273.15,  298.15,  323.15,  348.15,  373.15,  398.15))

volume = 0.044 # m**3
temps = np.array((273.15,  298.15,  323.15,  348.15,  373.15,  398.15))
pcyl_ig  = pig( volume, temps)
pcyl_rk  = prk( volume, temps, Tc=fluids['C4H10']['Tc'], Pc=fluids['C4H10']['Pc'])
pcyl_vdw = pvdw(volume, temps, Tc=fluids['C4H10']['Tc'], Pc=fluids['C4H10']['Pc'])

plt.plot(temps, pcyl_ig,  'g-',
         temps, pcyl_rk,  'r-',
         temps, pcyl_vdw, 'b-')
plt.show()
```

-   Plot the resulting error curves of `pcyl_ig` v. `pcyl_rk` and `pcyl_vdw` v. `pcyl_rk` across the range of temperature.  (There will be two curves in your plot.)

<!-- grade_id: plt-butane1, points: 1, grade: true, solution: true -->
```python
# YOU RECALCULATE THE VARIABLES HERE

# YOU WRITE THE plt.plot STATEMENT
plt.show()
```

Compared to argon or even oxygen, butane is pretty far away from ideality (the green line).

-   At the right-most data point (temperature $398.15\,\text{K}$), what is the percent difference between the ideal-gas law and Redlich–Kwong for butane, using $P_\text{rk}$ as the standard?  This is about 100 times worse of a match than argon to ideality.

<!-- grade_id: error-butane1, points: 1, grade: true, solution: true -->
```python
# WRITE ANY NECESSARY CODE HERE
# STORE YOUR ANSWER IN THE VARIABLE `ans`
```

Prior predictions were quite close to each other, within a few percent.  But for other cases they diverge dangerously (if one relies on the less-accurate equation).  Consider an amount of CO$_2$ gas compressed into the small volume of $0.000667\,\text{m}^3$ at a temperature of $248.15\,\text{K}$ ($-25\,^{\circ}\text{C}$).

-   What is the pressure according to the ideal-gas law?

<!-- grade_id: error-co21, points: 1, grade: true, solution: true -->
```python
# WRITE ANY NECESSARY CODE HERE
# STORE YOUR ANSWER IN THE VARIABLE `ans`
```

-   What is the pressure according to the Redlich-Kwong gas law?

<!-- grade_id: error-co22, points: 1, grade: true, solution: true -->
```python
# WRITE ANY NECESSARY CODE HERE
# STORE YOUR ANSWER IN THE VARIABLE `ans`
```

-   What is the pressure according to the van der Waals gas law?

<!-- grade_id: error-co23, points: 1, grade: true, solution: true -->
```python
# WRITE ANY NECESSARY CODE HERE
# STORE YOUR ANSWER IN THE VARIABLE `ans`
```

Clearly, at low temperatures and high pressures the laws start to break down, due to deviations from ideality and condensation.

Let's go even more extreme:  take a volume of 10% of that, or `0.0000667`.  (At this absurdly small volume, one can expect very high pressures to be necessary to contain the fluid, and our results become unphysical at some point---but that point is not determinable from the equations themselves!)

-   Plot the pressures of all three equations of state for CO$_{2}$ at across the temperature range $\{100, 125, 150, 175, 200\} \,^{\circ}\text{C}$ for $V = 0.0000667 \,\text{m}^3$.  (Note that you will need to convert from degrees centigrade to Kelvin, $T[K] = T[^{\circ}\text{C}] + 273.15$.)

<!-- grade_id: error-co2-comp, points: 1, grade: true, solution: true -->
```python
# YOU RECALCULATE THE VARIABLES HERE

# YOU WRITE THE plt.plot STATEMENT
plt.show()
```

-   At $T = 500\,\text{K}$, $V = 0.0000667 \,\text{m}^3$ for CO$_{2}$, what is the percent difference between the ideal-gas law and Redlich–Kwong, using $P_\text{rk}$ as the standard?

<!-- grade_id: compress1, points: 1, grade: true, solution: true -->
```python
# WRITE ANY NECESSARY CODE HERE
# STORE YOUR ANSWER IN THE VARIABLE `ans`
```

-   Under the same conditions for CO$_{2}$, what is the percent difference between the van der Waals law and Redlich–Kwong, using $P_\text{rk}$ as the standard?

<!-- grade_id: compress2, points: 1, grade: true, solution: true -->
```python
# WRITE ANY NECESSARY CODE HERE
# STORE YOUR ANSWER IN THE VARIABLE `ans`
```

-   Plot the pressures of all three equations of state for CO$_{2}$ across the temperature range $\{100, 110, 120, ..., 390, 400\} \,^{\circ}\text{C}$ for $V = 0.0000667 \,\text{m}^3$.  Use `np.linspace` to construct the range.

<!-- grade_id: temps-all, points: 1, grade: true, solution: true -->
```python
# range to use (for x-values)
temps = np.linspace(100,400,31)  # `linspace` conveniently breaks a range in N pieces---in this case, 31 divisions from 100 to 400

# YOU RECALCULATE THE VARIABLES HERE

# YOU WRITE THE plt.plot STATEMENT
plt.show()
```

Equations of state find broad application, but it's clear that you as an engineer need to understand the differences between them, the ranges of applicability, and the risks of under- or overestimating the pressure–temperature relationship in order to carry out sound engineering.